Dual-specific autophosphorylation of kinase IKK2 enables phosphorylation of substrate IκBα through a phosphoenzyme intermediate

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public Review):

The model of phosphotransfer from Y169 IKK to S32 IkBa is compelling and an important new contribution to the field. In fact, this model will not be without controversy, and publishing the work will catalyze follow-up studies for this kinase and others as well. As such, I am supportive of this paper, though I do also suggest some shortening and modification.

We appreciate the reviewers candid response on the difficulty of this study and the requirement of follow-up studies to confirm a direct transfer of the phosphate. We also have edited the manuscript to make it shorter.

Generally, the paper is well written, but several figures should be quantified, and experimental reproducibility is not always clear. The first 4 figures are slow-going and could be condensed to show the key points, so that the reader gets to Figures 6 and 7 which contain the "meat" of the paper.

We have indicated the experimental reproducibility in the methodology section against each assay. We have shortened the manuscript corresponding to sections describing figures 1-4. However, when we talked to some of our colleagues whose expertise do not align with kinases and IKK, we realized that some description were necessary to introduce them to the next figures. Additionally, we added Fig. S6 indicating that the radiolabelled phospho-IKK2 Y169F is unable to transfer its own phosphate group(s) to the substrate IkBa.

Reviewer #2 (Public Review):

Phosphorylation of IκBα is observed after ATP removal, although there are ambiguous requirements for ADP.

We agree with the reviewer that this observation is puzzling. We hypothesize that ADP is simultaneously regulating the transfer process likely through binding to the active site.

It seems that the analysis hinges on the fidelity of pan-specific phosphotyrosine antibodies.

We agree with the reviewer. To bolster our conclusion, we used antibodies from two different sources. These were Monoclonal mouse anti-Phospho-Tyrosine (catalogue number: 610000) was from BD Biosciences or from EMD Millipore (catalogue no. 05-321X).

The analysis often returns to the notion that tyrosine phosphorylation(s) (and critical active site Lys44) dictate IKK2 substrate specificity, but evidence for this seems diffuse and indirect. This is an especially difficult claim to make with in vitro assays, omitting the context of other cellular specificity determinants (e.g., localization, scaffolding, phosphatases).

We agree with the concerns that the specificity could be dependent on other cellular specificity determinants and toned down our claims where necessary. However, we would like to point out that the specificity of IKK2 towards S32 and S36 of IkBa in cells in response to specific stimuli is well-established. It is also well-established that its non-catalytic scaffolding partner NEMO is critical in selectively bringing IkBa to IKK from a large pool of proteins. The exact mechanism of how IKK2 choose the two serines amongst many others in the substrate is not clear.

Multiple phosphorylated tyrosines in IKK2 were apparently identified by mass spectrometric analyses, but the data and methods are not described. It is common to find non-physiological post-translational modifications in over-expressed proteins from recombinant sources. Are these IKK2 phosphotyrosines evident by MS in IKK2 immunoprecipitated from TNFa-stimulated cells? Identifying IKK2 phosphotyrosine sites from cells would be especially helpful in supporting the proposed model.

Mass spectrometric data for identification of phosphotyrosines from purified IKK2 is now incorporated (Figure S3A). Although we have not analyzed IKK2 from TNF-a treated cells in this study, a different study of phospho-status of cellular IKK2 indicated tyrosine phosphorylation (Meyer et al 2013).

Reviewer #3 (Public Review):

The identity and purity of the used proteins is not clear. Since the findings are so unexpected and potentially of wide-reaching interest - this is a weakness. Similar specific detection of phospho-Ser/Thr vs phospho-Tyr relies largely on antibodies which can have varying degrees of specificity.

We followed a stringent purification protocol of several steps (optimized for the successful crystallization of the IKK2) that removed most impurities (PMID: 23776406, PMID: 39227404). The samples analysed with ESI MS did not show any significant contaminating kinase from the Sf9 cells.

Sequence specific phospho-antibodies used in this study are very well characterized and have been used in the field for years (Basak et al 2007, PMID: 17254973). We agree on the reviewer’s concerns on the pan-specific phospho-antibodies. Since phospho-tyrosine detection is the crucial aspect of this study, we minimized such bias by using pan-specific phosphotyrosine antibodies from two independent sources.

Reviewer #1 (Recommendations For The Authors):

I understand that Figure 3 shows that K44M abolishes both S32/26 phosphorylation and tyrosine phosphorylation, but not PEST region phosphorylation. This suggests that autophosphorylation is reflective of its known specific biological role in signal transduction. But I do not understand why "these results strongly suggest that IKK2-autophosphorylation is critical for its substrate specificity". That statement would be supported by a mutant that no longer autophosphorylates, and as a result shows a loss of substrate specificity, i.e. phosphorylates non-specific residues more strongly. Is that the case? Maybe Darwech et al 2010 or Meyer et al 2013 showed this.

Later figures seem to address this point, so maybe this conclusion should be stated later in the paper.

We have now clarified this in the manuscript and moved the comment to the next section. We have consolidated the results in Figure 3 and 4 in the previous version into a single figure in Figure. The text has also been modified accordingly.

Page 10: mentions DFG+1 without a proper introduction. The Chen et al 2014 paper appears to inform the author's interest in Y169 phosphorylation, or is it just an additional interesting finding? Does this publication belong in the Introduction or the Discussion?

The position of Y169 at the DFG+1 was intriguing and the 2014 article Chen et al further bolstered our interest in this residue to be investigated. We think this publication is important in both sections.

To understand the significance of Figure 4D, we need a WT IKK2 control: or is there prior literature to cite? This is relevant to the conclusion that Y169 phosphorylation is particularly important for S32 phosphorylation.

We have now added a new supplementary figure where activities of WT and Y169F IKK2 towards WT and S32/S36 mutants are compared (Figure S3F). At a similar concentration, the activity of WT-IKK2 is many fold higher than that of YtoF mutants (Fig. 4C). The experiments were performed simultaneously, although samples were loaded on different gels but otherwise processed in a similar way. The corresponding data is now included in the manuscript as Figure S3F.

The cold ATP quenching experiment is nice for testing the model that Y169 functions as a phospho sink that allows for a transfer reaction. However, there is only a single timepoint and condition, which does not allow for a quantitative analysis. Furthermore, a positive control would make this experiment more compelling, and Y169F mutant should show that cold ATP quenching reduces the phosphorylation of IkBa.

We thank the reviewer for appreciating our experimental design, and pointing out the concerns. We kept the ATP-time point as the maximum of the non-competition experiment. Also, we took 50mM ATP to compare its competition with highest concentration of ADP used. The idea behind using the maximum time and ATP (comparable to ADP) was to capture the effect of competitive-effect of ATP, if any, that would be maximal in the given assay condition in comparison with the phospho-transfer set up in absence of cold ATP. We agree that finer ranges of ATP concentration and time points would have enabled more quantitative analyses. We have now included data where different time intervals are tested (Figure S5D).

Why is the EE mutant recognized by anti-phospho-serine antibodies? In Figure 2F.

We anticipate Serine residues besides those in the activation loop to be phosphorylated when IKK2 is overexpressed and purified from the Sf9 cells. Since Glu (E) mimics phospho-Ser, the said antibody cross reacts with the IKK2-EE that mimics IKK2 phosphorylated at Ser177 and 181.

Figure 7B is clear, but 7C does not add much.

We have now removed the Fig. 7C in the current version. Figure 7 is now renumbered as Figure 6 that does not contain the said cartoon.

Reviewer #2 (Recommendations For The Authors):

Regarding the specificity arguments (see above in public review), the authors note that NEMO is very important in IKK specificity, and - if I'm understanding correctly - most of these assays were performed without NEMO. Would the IKK2-NEMO complex change these conclusions?

NEMO is a scaffolding protein whose action goes beyond the activation of the IKK-complex. In cells, NEMO brings IkBa from a pool of thousands of proteins to its bonafide kinase when the cells encounter specific signals. In other words, NEMO channels IKK-activity towards its bonafide substrate IkBa at that moment. Though direct proof is lacking, it is likely that NEMO present IkBa in the correct pose to IKK such that the S32/S36 region of IkBa is poised for phosphorylation. The proposed mechanism in the current study further ensures the specificity and fidelity of that phosphorylation event. We believe this mechanism will be preserved in the IKK-NEMO complex unless proven otherwise. As shown below, IKK2 undergoes tyrosine autophosphorylation in presence of NEMO.

The work primarily focuses on Y169 as a candidate target for IKK autophosphorylation. This seems reasonable given the proximity to the ATP gamma phosphate. However, Y188F more potently disrupted IκBα phosphorylation. The authors note that this could be due to folding perturbations, but this caveat would also apply to Y169F. A test for global fold perturbations for both Tyr mutants would be helpful.

Y188 is conserved in S/T kinases and that in PKA (Y204) has been studied extensively using structural, biochemical and biophysical tools. It was found in case of PKA that Y204 participates in packing of the hydrophobic core of the large lobe. Disruption of this core structure by mutation allosterically affect the activity of the kinase. We also observed similar engagement of Y188 in IKK2’s large lobe, and speculated folding perturbations in analogy with the experimental evidence observed in PKA. What we meant was mutation of Y188 would allosterically affect the kinase activity. Y169 on the other hand is unique at that position, an no experimental evidence on the effect of phospho-ablative mutation of this residue exist in the literature. Hence, we refrained from speculating its effect on the folding or conformational allostery, however, such a possibility cannot be ruled out.

I struggled to follow the rationalization of the results of Figure 4D, the series of phosphorylation tests of Y169F against IκBα with combinations of phosphoablative or phosphomimetic variants at Ser32 and Ser36. This experiment is hard to interpret without a direct comparison to WT IKK2.

We agree with the reviewer’s concerns. Through this experiment we wanted to inform about the importance of Tyr-phosphorylation of IKK2 in phosphorylating S32 of IκBα which is of vital importance in NF-kB signaling. We have now provided a comparison with WT-IKK2 in the supplementary Figure S3F. We hope this will help bring more clarity to the issue.

MD simulations were performed to compare structures of unphosphorylated vs. Ser-phosphorylated (p-IKK2) vs. Ser+Tyr-phosphorylated (P-IKK2) forms of IKK2. These simulations were performed without ATP bound, and then a representative pose was subject to ADP or ATP docking. The authors note distortions in the simulated P-IKK2 kinase fold and clashes with ATP docking. Given the high cellular concentration of ATP, it seems more logical to approach the MD with the assumption of nucleotide availability. Most kinase domains are highly dynamic in the absence of substrate. Is it possible that the P-IKK2 poses are a result of simulation in a non-physiological absence of bound ATP? Ultimately, this MD observation is linked to the proposed model where ADP-binding is required for efficient phospho-relay to IκBα. Therefore, this observation warrants scrutiny. Perhaps the authors could follow up with binding experiments to directly test whether P-IKK2 binds ADP and fails to bind ATP.

We thank that reviewer for bringing up this issue. This is an important issue and we must agree that we don’t fully understand it yet. We took more rigorous approach this time where we used three docking programs: ATP and ADP were docked to the kinase structures using LeDock and GOLD followed by rescoring with AutoDock Vina. We found that ATP is highly unfavourable to P-IKK2 compared to ADP. To further address these issues, we performed detailed MM-PBSA (Molecular Mechanics Poisson-Boltzmann Surface Area) analyses after MD-simulation to estimate binding free energies and affinities of ADP and ATP for each of the three differently phosphorylated states of IKK2. These analyses (Figure S4 E and F) clearly indicate that phosphorylated IKK2 have much higher preference for ADP over ATP. However, it does not negate ATP-binding by P-IKK2 in a different pose that may not support kinase activity.

We could not perform any binding experiment because of the following reason. We incubated FL IKK2 WT with or without cold ATP for 30mins, and then incubated these samples with 32P-ATP and analysed the samples by autoradiography after resolving them on a 10% SDS-PAGE. We found that even after pre-incubation of the kinase with excess cold ATP it still underwent autophosphorylation when radioactive ATP was added as shown below. This prevented us from doing direct binding experiment with ATP as it would not represent true binding event. We also noticed that after removal of bulk ATP post autophosphorylation, phosphorylated IKK2 is capable of further autophosphorylation when freshly incubated with ATP. We have not been able to come up with a condition that would only account for binding of ATP and not hydrolysis.

The authors could comment on whether robust phosphorylation of NEMO was expected (Figure 1D). On a related note, why is NEMO a single band in the 1D left panel and double bands on the right?

No, we did not expect robust phosphorylation of NEMO. However, robust phosphorylation of NEMO is observed only in the absence of IκBα. In presence of IκBα, phosphorylation of NEMO goes down drastically. These were two different preparations of NEMO. When TEV-digestion to remove His-tag is incomplete it gives two bands as the tagged and untagged versions cannot be separated in size exclusion chromatography which is the final step.

Page 14, line 360. "...observed phosphorylation of tyrosine residue(s) only upon fresh ATP-treatment..." I'm not sure I understand the wording here (or the relevance of the citation). Is this a comment on unreported data demonstrating the rapid hydrolysis of the putative phosphotyrosine(s)? If so, that would be helpful to clarify and report in the supporting information.

In our X-ray crystallographic studies with phosphorylated IKK2 we failed to observe any density of phosphate moiety. Furthermore, this IKK2 showed further autophosphorylation when incubated with fresh ATP. These two observations lead us to believe that some of the autophosphorylation are transient in nature. However, quantitative kinetic analyses of this dephosphorylation have not been performed.

Figure S3 middle panel: The PKA substrate overlaid on the IKK2 seems sterically implausible for protein substrate docking. Is that just a consequence of the viewing angle? On a related note, Figure S3 may be mislabeled as S4 in the main text).

It is a consequence of the viewing angle. Also, we apologize for this inadvertent mislabelling. It has been corrected in the current version.

Reviewer #3 (Recommendations For The Authors):

The detection of phosphorylated amino acids relies largely on antibodies which can have a varying degree of specificity. An alternative detection mode of the phospho-amino acids for example by MS would strengthen the evidence.

We agree with the concern of specificity bias of antibodies. We tried to minimize such bias by using two different p-Tyr antibodies as noted previously and also in the methodology section. We were also able to detect phospho-tyrosine residues by MS/MS analyses, representative spectra are now added (Figure S3A).

IKK2 purity - protocol states "desired purity". What was the actual purity and how was it checked? MS would be useful to check for the presence of other kinases.

Purity of the recombinantly purified IKK2s are routinely checked by silver staining. A representative silver stained SDS-PAGE is shown (Figure S1C). It may be noted that, there’s a direct correlation of expression level and solubility, and hence purification yield and quality with the activity of the kinase. Active IKK2s express at much higher level and yields cleaner prep. In our experience, inactive IKKs like K44M give rise to poor yield and purity. We analysed K44M by LC MS/MS to identify other proteins present in the sample. We did not find any significant contaminant kinase the sample (Figure S1D). The MS/MS result is attached.

Figure 1C&D: where are the Mw markers? What is the size of the band? What is the MS evidence for tyrosine phosphorylation?

We have now indicated MW marker positions on these figures.

MS/MS scan data for the two peptides containing pTyr169 and pTyr188 are shown separately (Figure S3A).

Figure 2F: Why is fresh ATP necessary? Why was Tyr not already phosphorylated? The kinetics of this process appear to be unusual when the reaction runs to completion within 5 minutes ?

As stated earlier, we believe some of the autophosphorylation are transient in nature. We think the Tyr-phosphorylation are lost due to the action of cellular phosphatases. We agree with the concern of the reviewer that, the reaction appears to reach completion within 5 minutes in Fig 2F. We believe it is probably due to the fact that the amount of kinase used in this study exceeds the linear portion of the dynamic range of the antibody used. Lower concentration of the kinase do show that reaction does not reach completion until 60mins as shown in Fig. 2A.

Figure 3: Can the authors exclude contamination with a Tyr kinase in the IKK2-K44M prep? The LC/MS/MS data should be included.

We have reanalysed the sample on orbitrap to check if there’s any Tyr-kinase or any other kinase contamination. We used Spodoptera frugiperda proteome available on the Uniprot website for this analysis. These analyses confirmed that there’s no significant kinase contaminant present in the fraction (Figure S1D).

What is the specificity of IKK-2 Inhibitor VII? Could it inhibit a contaminant kinase?

This inhibitor is highly potent against IKK2 and the IKK-complex, and to a lesser extent to IKK1. No literature is available on its activity on other kinases. In an unrelated study, this compound was used alongside MAPK inhibitor SB202190 wherein they observed completely different outcomes of these two inhibitors (Matou-Nasri S, Najdi M, AlSaud NA, Alhaidan Y, Al-Eidi H, Alatar G, et al. (2022) Blockade of p38 MAPK overcomes AML stem cell line KG1a resistance to 5-Fluorouridine and the impact on miRNA profiling. PLoS ONE 17(5):e0267855. https://doi.org/10.1371/journal.pone.0267855). This study indirectly proves that IKK inhibitor VII does not fiddle with the MAPK pathways. We have not found any literature on the non-specific activity of this inhibitor.

Figure 6B: the band corresponding to "p-IkBa" appears to be similar in the presence of ADP (lanes 4-7) or in the absence of ADP but the presence of ATP (lane 8).

Radioactive p-IκBα level is more when ADP is added than in absence of ADP. In presence of cold ATP, radioactive p-IκBα level remains unchanged. This result strongly indicate that the addition of phosphate group to IκBα happens directly from the radioactively labelled kinase that is not competed out by the cold ATP.

Changes from the earlier version to the revised version

Abstract

Abstract

Rapid and high-fidelity phosphorylation of two serines (S32 and S36) of IκBα by a prototype Ser/Thr kinase IKK2 is critical for fruitful canonical NF-κB activation. Here, we report that IKK2 is a dual specificity Ser/Thr kinase that autophosphorylates itself at tyrosine residues in addition to its activation loop serines. Mutation of one such tyrosine, Y169, located in proximity to the active site, to phenylalanine, renders IKK2 inactive for phosphorylation of S32 of IκBα. Surprisingly, auto-phosphorylated IKK2 relayed phosphate group(s) to IκBα without ATP when ADP is present. We also observed that mutation of K44, an ATP-binding lysine conserved in all protein kinases, to methionine renders IKK2 inactive towards specific phosphorylation of S32 or S36 of IκBα, but not non-specific substrates. These observations highlight an unusual evolution of IKK2, in which autophosphorylation of tyrosine(s) in the activation loop and the invariant ATP-binding K44 residue define its signal-responsive substrate specificity ensuring the fidelity of NF-κB activation.

Introduction

Protein kinases bestow novel identity to their substrates by adding phosphate moieties at the site of modification that often play key regulatory roles in a diverse array of signaling events. For example, inflammatory or pathogenic signals trigger precisely regulated induction of the transcription factor NF-κB in metazoans(Hoffmann and Baltimore, 2006) that depends upon the upstream activation of the Inhibitor of κB Kinase-complex (IKK-complex), comprising two catalytic subunits, IKK1 (also known as IKKα) and IKK2 (or IKKβ), and an essential regulatory scaffolding protein, NEMO (IKKγ)(DiDonato et al., 1997; Rothwarf et al., 1998; Zandi et al., 1997). IKK2 activity is maintained at a very low basal level in resting cells(Ghosh and Karin, 2002; Hacker and Karin, 2006; Hinz and Scheidereit, 2014; Liu et al., 2012). Upon activation, IKK2 phosphorylates the N-terminus of IκBα (Inhibitor of κBα) protein specifically at two Serine residues (S32 and S36) marking it for ubiquitinylation-mediated proteasomal degradation (Figure 1A). Signal-responsive phosphorylation of IκBα refers to the phosphorylation of S32 and S36 by the IKK-activity. Degradation of IκBα liberates NF-κB to execute its gene expression program (Hayden and Ghosh, 2008; Hinz and Scheidereit, 2014; Karin and Ben-Neriah, 2000; Scheidereit, 2006). Mutation of both S32 and S36 to alanine residues converts IκBα into a non-phosphorylatable super-repressor of NF-κB(Brown et al., 1995; Lin et al., 1995). Furthermore, IKK possesses exquisite specificity for these two serines, and it fails to phosphorylate IκBα when S32 and S36 are mutated even to another phosphorylatable residue, threonine(DiDonato et al., 1997).

Autophosphorylation of IKK2 at hitherto uncharacterized sites.

(A) NEMO is known to remain bound to the IKK-complex pre-and post-stimulation, helps in activation as well as substrate recognition to channelize IKK-activity towards IκBα. However, it is not clear how IKK specifically phosphorylates the Ser32&36 of IκBα. (B) Domain organisation of FL IKK2 depicting functional kinase domain (KD), ubiquitin-like domain (ULD), and scaffold dimerization domain (SDD) based on the X-ray structures. Known sites of Ser phosphorylation and NEMO interacting regions at the extreme C-termini are highlighted. Serine and tyrosine residues in the activation loop are shown, mutation of which to Glu renders IKK2 constitutively active. The conserved ATP-interacting Lys (at position 44) is also shown. (C) Autophosphorylation of FL IKK2WT upon incubation with radioactive ATP for different time periods in an in vitro kinase assay. (D) In vitro radioactive kinase assays performed to check the effect of NEMO on autophosphorylation of FL IKK2WT (left panel), and that of NEMO and IκBα on autophosphorylation and substrate phosphorylation (right panel) activity of FL IKK2WT. (E) Effect of Inhibitor VII on FL IKK2WT autophosphorylation and IκBα substrate phosphorylation at different inhibitor concentrations in a radioactive in vitro kinase assay. A schematic representation of the assay procedure is shown in Figure S1B. This assay was performed twice. (F) Radioactive kinase assay conducted to check auto-and substrate phosphorylation of full length and deletion constructs of constitutively active form of IKK2 harbouring phosphomimetic Ser177Glu and Ser181Glu mutations.

IKK2 is an essential Ser/Thr kinase that ensures the activation of a particular signaling cascade -NF-κB, with high fidelity despite its reported pleiotropy in other contexts(Antonia et al., 2021; Schröfelbauer et al., 2012; Schrofelbauer and Hoffmann, 2011). IKK2 is a multidomain protein consisting of a kinase domain (KD) followed by a Ubiquitin-like domain (ULD) and the scaffold dimerization domain (SDD) (Figure 1B). IKK2 toggles between inactive and active states based on the phosphorylation status of two serine residues (S177 and S181; S176 and S180 in the case of IKK1) within its activation loop (activation loop) in the KD(Huse and Kuriyan, 2002; Liu et al., 2013; Polley et al., 2013; Xu et al., 2011). Replacement of these two serines with phospho-mimetic glutamate residues (S177E, S181E; henceforth EE) renders the kinase constitutively active. IKK2 is reported to be primed through phosphorylation at S177 by the upstream kinase TAK1 that leads to auto-phosphorylation at S181, thus rendering the kinase fully active in cells(Zhang et al., 2014). However, oligomerization upon association with NEMO, in the presence of linear or Lys-63 linked poly-ubiquitin (Ub) chains or at a high concentration of IKK2, also enables trans auto-phosphorylation at S177 and S181(Chen, 2012; Du et al., 2022; Ea et al., 2006; Polley et al., 2013). It was also shown that NEMO not only helps in the activation of the IKK-complex, but also ensures that the IKK-activity is well directed towards IκBα in the inhibitor bound transcriptionally inactive NF-κB (NF-κB: IκBα complex)(Schröfelbauer et al., 2012). Activated IKK2 phosphorylates IκBα specifically at residues S32/S36, for which the C-terminal SDD of IKK2 (residue range 645-756) is critical – a shorter construct of IKK2 lacking the NBD and SDD displayed CK2-like activity that did not retain the exquisite specificity of IKK2 to phosphorylate the S32/36 of IκBα (Shaul et al., 2008). Other sites of IκBα e.g., the residues in the C-terminal PEST region that contains proline (P), glutamic acid (E), serine (S) and threonine (T) residues, but not S32/S36, can be phosphorylated by various kinases(Barroga et al., 1995; Tergaonkar et al., 2003). The role of IκBα-phosphorylation at sites other than the signal responsive S32/S36 phosphorylation in NF-κB activation and signaling is elusive. These findings underscore the ability of IKK2 to function specifically in order to guarantee rapid NF-κB activation in canonical signaling. It is however, still not fully understood how IKK2 achieves its exquisite specificity towards S32/S36 of IκBα to mark it for proteasomal degradation that liberates active NF-κB dimers.

Here we report an in-depth analysis of the catalytic features of IKK2 that revealed that IKK2 undergoes multisite autophosphorylation and can phosphorylate its own tyrosine residue(s) in addition to its own serines or the serines of its substrate IκBα, i.e., IKK2 possesses dual specificity. This phosphorylation at tyrosine residue(s) is autocatalytic. This phosphorylation at tyrosine residues is dependent upon the prior activation of IKK2 via activation loop S177/S181 phosphorylation, and is critical to phosphorylation of signal responsive S32/S36 of IκBα. Mutation of Y169 of IKK2 to phospho-ablative phenylalanine severely compromised phosphorylation of IκBα at S32. Mutation of Y188 to phenylalanine rendered the kinase practically inactive in vitro. IKK2 activity towards signal responsive S32/S36 of IκBα was severely compromised in MEF cells reconstituted with these two mutant IKK2s in response to treatment with TNFα. We also observed that mutation of the conserved ATP-interacting residue K44 in IKK2 to methionine leads to a loss of tyrosine as well as S177,S181 auto-phosphorylation activity of the kinase along with its ability to phosphorylate S32/S36 of IκBα, thereby clearly suggesting a critical role of this Lysine residue in regulating IKK2’s kinase activity for tyrosine as well as serine phosphorylations. Interestingly, IKK2 K44M retained non-specific kinase activity referring to phosphorylation of IκBα at residues other than S32 and S36. Furthermore, we observed that, after auto-phosphorylation at tyrosine, IKK2 can phosphorylate IκBα in the absence of exogenous supply of ATP, as long as ADP is present in the reaction. Overall, our results reveal unique, and likely transient, autophosphorylated intermediates of IKK2 that appear to be critical in relaying phosphoryl group(s) specifically to S32/S36 of IκBα rather than following the conventional mode of phosphate group transfer directly from ATP to the substrate seen in eukaryotic protein kinases. IKK2 thus ensures fidelity of NF-κB activation. Recent reports in other kinases indicate that multisite (auto)phosphorylation is a major determinant in regulating function of respective kinases and regulation exerted by those kinases. Our study adds another layer of regulation and critical nature of such autophosphorylation that plays a significant biological role in metazoans.

Results

IKK2 undergoes auto-phosphorylation at uncharacterized sites

Earlier version: Signal-induced phosphorylation of activation loop Ser177/Ser181 marks the general activation of IKK2 (Figure 1B) from an unphosphorylated inactive state. Phosphorylation of several other serines located within the flexible C-terminal region of IKK2, spanning amino acid 701 to the end, has also been reported (Figure 1B)(Delhase et al., 1999). Hyperphosphorylation of these C-terminal sites down-regulates IKK2 activity in cells. It was also reported earlier that, ectopically overexpressed IKK2 is capable of autophosphorylating these two serines in the activation loop in trans even in the absence of an activating physiological signal. Indeed, recombinantly purified FL IKK2WT is capable of autophosphorylation in vitro when incubated with γ32P-ATP that increases with the time of incubation (Figure 1C). The spread-out of slower migrating phospho-IKK2 band likely indicates phosphorylation of the protein at multiple sites. We also observed that autophosphorylation of IKK2 occurred irrespective of the presence of its cognate binding partner NEMO and bonafide substrate IκBα (Figure 1D). To confirm that this phosphorylation of IKK2 was indeed autophosphorylation and not due to an unrelated kinase copurified with IKK2, we generated a K44M mutant of full-length IKK2 where the ATP-binding invariant lysine residue conserved in Ser/Thr/Tyr kinases was mutated to methionine. This lysine residue anchors ATP in the respective binding pocket of the kinase by forming hydrogen bonds with the α- and β-phosphate of ATP that is critical for productive phosphorylation reaction in protein kinases. Mutating this K to M is expected to greatly reduce or eliminate the kinase activity. An in vitro kinase assay using this mutant along with the WT kinase using γ32P-ATP as the phosphor-donor confirmed that the IKK2 K44M lacked both IκBα (substrate) and autophosphorylation activities (Figure S1A). That these phosphorylations were indeed autophosphorylation was further supported by another in vitro radioactive kinase assay performed in the presence or absence of an IKK-specific ATP-competitive inhibitor, Calbiochem Inhibitor VII (Figure 1E) where the kinase and kinase-substrate reactions were incubated for 30 minutes with the inhibitor prior to adding the phosphor-donor, γ32P-ATP. A similar assay was performed with cold ATP and specific phosphorylation of S32/36 on IκBα was monitored by immunoblotting in presence and absence of the same inhibitor (Figure S1B). Purification status of recombinant FL IKK2 WT and FL IKK2 K44M proteins used in this study are shown in Figure S1C.

Revised version: Signal-induced phosphorylation of activation loop Ser177/Ser181 marks the general activation of IKK2 (Figure 1B). Hyperphosphorylation of several other serines located within the flexible C-terminal region of IKK2, spanning amino acid 701 to the end (Figure 1B), has been reported to down-regulate IKK2 activity in cells (Delhase et al., 1999). Furthermore, ectopically overexpressed IKK2 was found to be capable of autophosphorylating these two serines in the activation loop in trans even in the absence of an activating physiological signal {Polley, 2013 #27}. Indeed, recombinantly purified FL IKK2WT is capable of autophosphorylation in vitro when incubated with γ32P-ATP that increases with the time of incubation (Figure 1C), possibly at multiple sites as indicated by the spread-out of slower migrating phospho-IKK2 band. We also observed that autophosphorylation of IKK2 occurred irrespective of the presence of its cognate binding partner NEMO and bonafide substrate IκBα (Figure 1D). An in vitro kinase assay using a kinase inactivating mutant, IKK2 K44M along with the WT kinase using γ32P-ATP as the phosphor-donor confirmed that the above-mentioned phosphorylation of IKK2 was indeed autophosphorylation and not due to an unrelated kinase copurified with IKK2 as IKK2 K44M lacked both IκBα (substrate) and autophosphorylation activities (Figure S1A). Lys44 is the ATP-binding invariant lysine residue in IKK2 that is conserved in STYKs, mutation to which to Met is expected to greatly reduce or eliminate the kinase activity. These observations were further supported by another in vitro radioactive kinase assay performed in the presence or absence of an IKK-specific ATP-competitive inhibitor, Calbiochem Inhibitor VII (Figure 1E), that inhibits specific phosphorylation of S32/36 on IκBα (Figure S1B). In both cases kinase and kinase-substrate reactions were incubated for 30 minutes with the inhibitor prior to adding the phosphor-donor, γ32P-ATP or cold ATP. Purification status of recombinant FL IKK2 WT and FL IKK2 K44M proteins used in this study are shown in Figure S1C, along with LC MS/MS study performed on the FL IKK2 K44M (Figure S1D).

Removed in the revised version: Next, we used different versions of the IKK2 protein where all previously known phosphorylation sites were lacking: Ser177 and 181 were mutated to phosphor-mimetic Glu (EE; this version of IKK2 is known to be constitutively active (Zandi et al., 1998)), and the C-terminal serine-rich region was truncated to different extents. The shortest construct,

Earlier version: Δ664EE that entirely lacked the flexible C-terminus (670-756) was further subjected to limited proteolysis using trypsin to eliminate residues, if any, from the flexible ends of the enzymes and this trypsin-treated IKK2 was purified prior to in vitro kinase assays. Each of these constructs was active in phosphorylating itself and the substrate IκBα (Figure 1F). The occurrence of autophosphorylation despite the absence of activation loop serines and serines within the C-terminus strongly suggested the presence of hitherto uncharacterized sites of autophosphorylation in IKK2.

Revised version: Next, we used different versions of the IKK2 protein where all previously known phosphorylation sites were lacking: Ser177 and 181 were mutated to phosphor-mimetic Glu (EE; this version of IKK2 is known to be constitutively active (Zandi et al., 1998)), and the C-terminal serine-rich region was truncated to different extents. The shortest construct, Δ664EE that entirely lacked the flexible C-terminus (670-756) was further subjected to limited proteolysis using trypsin to eliminate residues, if any, from the flexible ends of the enzymes and this trypsin-treated IKK2 was purified prior to in vitro kinase assays. Each of these constructs was active in phosphorylating itself and the substrate IκBα (Figure 1F). The occurrence of autophosphorylation despite the absence of activation loop serines and serines within the C-terminus strongly suggested the presence of hitherto uncharacterized sites of autophosphorylation in IKK2.

Auto-phosphorylation of IKK2 reveals its dual specificity

Earlier version: IKK2 is known as a bonafide Ser/Thr kinase that undergoes autophosphorylation. In the previous section, we found that IKK2’s autophosphorylation activity might extend beyond the previously characterized site(s). We asked if it is a dual specificity kinase. To test if Tyr residues can undergo auto-phosphorylation in addition to Ser, IKK2 incubated with Mg2+-ATP for various time periods was probed with antibodies that specifically recognize activation loop serines of IKK2 (Figure 2A, upper panel) or phospho-tyrosine residues (Figure 2A, middle panel). Unexpectedly, we observed de novo auto-phosphorylation of IKK2 on tyrosine residues along with the anticipated auto-phosphorylation of activation loop serines. It may be noted that autophosphorylation activation loop-S177/181 also increased with time, indicating inhomogeneous and incomplete phosphorylation of activation loop-S177/181 in the recombinant IKK2 when purified from Sf9 cells. We tested another phospho-tyrosine specific antibody to further substantiate this de novo tyrosine autophosphorylation of IKK2. It confirmed that detection of de novo phosphotyrosine on IKK2 was independent of the source of the antibody used (Figure 2B). The tyrosine auto-phosphorylation was abolished in the presence of an IKK-specific inhibitor, Inhibitor VII (Figure 2C). A similar assay was performed to monitor the phospho-tyrosine level in IKK2 as a function of urea concentration. This revealed the detrimental effects of urea on tyrosine autophosphorylation of IKK2 at concentrations beyond 1M (Figure S2A). A more sensitive assay with γ32P-ATP as the phospho-donor also indicated that urea concentrations beyond 1M inhibited both auto-phosphorylation (pan-phosphorylation) and substrate (IκBα) phosphorylation (Figure S2B). Furthermore, the IKK2-K44M mutant (kinase inactivating) did not display any tyrosine autophosphorylation when incubated with ATP unlike WT IKK2 (Figure 2D). These results confirm the novel tyrosine phosphorylation ability of IKK2, demonstrating its dual specificity to both its own Ser and Tyr residues. We also observed that activity of IKK2 towards S32 and S36 of IκBα and its own tyrosine residue(s) were affected similarly when the kinase assays were performed in the presence of various kinase inhibitors, both general (AMPPNP and Staurosporine) and IKK-specific (MLN-120B, TPCA, Inhibitor VII) (Figure S2C). These observations confirm that canonical activity of IKK2 towards S32 and S36 of IκBα, and towards its own tyrosine residues are mediated by the same active site.

Revised version: IKK2 is known as a bonafide Ser/Thr kinase that undergoes autophosphorylation. In the previous section, we found that IKK2’s autophosphorylation activity might extend beyond the previously characterized site(s). We asked if it is a dual specificity kinase. To this end, we unexpectedly observed de novo auto-phosphorylation of IKK2 on tyrosine residues (Figure 2A, middle panel) along with the anticipated auto-phosphorylation of activation loop serines (Figure 2A, upper panel) when IKK2 was incubated with Mg2+-ATP for various time periods. Increase in autophosphorylation of activation loop-S177/181 with time indicates inhomogeneous and incomplete phosphorylation of activation loop-S177/181 in the recombinant IKK2 when purified from Sf9 cells. To circumvent the specificity-bias of the pan-phosphotyrosine antibody used in Figure 2A, we used another phospho-tyrosine specific antibody, which confirmed that detection of de novo phosphotyrosine on IKK2 was independent of the source of the antibody used (Figure 2B). The tyrosine auto-phosphorylation was abolished in the presence of an IKK-specific inhibitor, Inhibitor VII (Figure 2C), and the IKK2-K44M mutant (kinase inactivating) did not display any tyrosine autophosphorylation when incubated with ATP unlike WT IKK2 (Figure 2D). IKK2 autophosphorylation assay performed in presence of different concentration of Urea revealed that its tyrosine-autophosphorylation was abolished at urea concentration beyond 1M (Figure S2A). A more sensitive assay with γ32P-ATP as the phospho-donor also indicated that urea concentrations above 1M abolished both pan-auto-phosphorylation and substrate (IκBα) phosphorylation of IKK2 (Figure S2B). These results confirmed that IKK2 is a dual specificity kinase that not only phosphorylates Ser on substrates, but also phosphorylates its own Ser and Tyr residues. The fact that activities of IKK2 towards S32 and S36 of IκBα and its own tyrosine residue(s) were affected similarly by various kinase inhibitors, both general (AMPPNP and Staurosporine) and IKK-specific (MLN-120B, TPCA, Inhibitor VII) (Figure S2C), confirm that both substrate phosphorylation and autophosphorylation activities of IKK2 are carried out by the same active site. We further observed that the C-terminal NEMO binding domain (NBD) was not required for autocatalytic dual specificity of IKK2 as various truncated versions of IKK2 (lacking NBD) were capable of phosphorylating itself on tyrosine residue(s), and phosphorylating IκBα at signal-responsive S32 and S36 (Figure 2E). Assessment of phospho-Ser and Tyr on a constitutively active NBD-deficient IKK2 (IKK2Δ664EE) monitored by monoclonal phospho-serine and phospho-tyrosine antibodies revealed de novo tyrosine autophosphorylation of the kinase while de novo serine autophosphorylation (if any) was not discernible in this assay condition (Figure 2F). The phospho-ablative S177A, S181A (AA; non-phosphorylatable activation loop) mutant of IKK2 that was incapable of phosphorylating S32/36 of IκBα (Figure S2D), also lacked tyrosine autophosphorylation activity (Figure 2G) in in vitro kinase assay suggesting that activation loop S177/S181 serine phosphorylation is essential in conducting dual specificity of IKK2.

IKK2 possesses autocatalytic dual specificity.

Earlier version: (A) In vitro kinase assay, followed by immunoblotting with FL IKK2WT using antibodies against phosphor-Ser (177/181) and phospho-Tyr. (B) A different commercial source of phospho-Tyr antibody was used to detect the same. (C) Effect of Inhibitor VII on tyrosine autophosphorylation of FL IKK2WT. (D) Immunoblotting performed with phospho-Tyr antibody to check autophosphorylation of IKK2 mutant K44M at different time points compared to that of IKK2 WT. (E) Tyrosine autophosphorylation and substrate phosphorylation of GST (1-54) IκBα in presence of full length and deletion mutants of IKK2 harbouring phosphomimetic Ser177Glu and Ser181Glu mutations. (F) De novo auto-phosphorylation of IKK2 on tyrosine monitored by phospho-Ser and phospho-Tyr specific monoclonal antibodies. (G) Autophosphorylation and subsequent immunoblot analysis of FLIKK2 mutant S177A, S181A using antibodies against phospho-IKK2-Ser(177/181) and phospho-Tyr.

Revised version: (A) In vitro kinase assay, followed by immunoblotting with FL IKK2WT using antibodies against phosphor-Ser (177/181) and phospho-Tyr. (B) A different commercial source of phospho-Tyr antibody was used to detect pTyr on IKK2. (C) Effect of Inhibitor VII on tyrosine autophosphorylation of FL IKK2WT. (D) Immunoblotting performed with phospho-Tyr antibody to check autophosphorylation of IKK2 mutant K44M at different time points compared to that of IKK2 WT. (E) Tyrosine autophosphorylation and substrate phosphorylation of GST (1-54) IκBα in presence of full length and deletion mutants of IKK2 harbouring phosphomimetic Ser177Glu and Ser181Glu mutations. (F) De novo auto-phosphorylation of IKK2 on tyrosine monitored by phospho-Ser and phospho-Tyr specific monoclonal antibodies. (G) Autophosphorylation and subsequent immunoblot analysis of FLIKK2 mutant S177A, S181A using antibodies against phospho-IKK2-Ser(177/181) and phospho-Tyr.

Removed in the revised version: Next, we checked the requirement of the C-terminal NEMO binding domain (NBD) for autocatalytic dual specificity. Constructs of various truncations (lacking NBD) were capable of phosphorylating itself on tyrosine residue(s), and phosphorylating IκBα at signal-responsive S32 and S36 (Figure 2E), indicating that the NBD domain is not required for either its ability to phosphorylate S32 and S36, as well as autophosphorylating itself on tyrosine. Next, we assessed the autocatalysis at Ser and Tyr by IKK2 using a constitutively active NBD-truncated version of IKK2 (IKK2Δ664EE). As anticipated, the protein was recognized by monoclonal phospho-serine antibody with or without fresh ATP-treatment; however, phospho-tyrosine antibody recognized it strongly only upon fresh ATP-treatment (Figure 2F) further confirming dispensability of the NBD in de novo tyrosine autophosphorylation of IKK2.

Removed in the revised version: Next, we asked if phosphorylation of activation loop S177/S181 was necessary for this dual specificity of IKK2. To this end, we purified the S177A, S181A (AA; non-phosphorylatable activation loop) phospho-ablative mutant of IKK2, and checked its tyrosine autophosphorylation activity. We found that the IKK2AA mutant lacked tyrosine autophosphorylation activity (Figure 2G) in an in vitro kinase assay. It suggests that activation loop S177/S181 serine phosphorylation is essential in conducting dual specificity in IKK2. Indeed, the IKK2 AA mutant was incapable of phosphorylating S32/36 of IκBα as well (Figure S2D) in an in vitro kinase assay.

Removed in the revised version: IKK2 autophosphorylation is critical to specificity of IκBα phosphorylation

Removed in the revised version: As mentioned earlier, it was long known that activation loop-serine phosphorylation was necessary for IKK-activation; we have now established that IKK2 phosphorylates itself on tyrosine residue(s) in addition to activation loop-serines. We wanted to further understand the significance of IKK2 autophosphorylation in the context of IκBα S32/S36 phosphorylation. It is known that disruption of a universal salt bridge mediated by the ATP-binding Lys (K44 in IKK2) is detrimental for kinase activity. However, it is not known if disruption of this salt bridge has any bearing on substrate specificity of eukaryotic protein kinases (EPKs). During our investigation, we noted unusual behavior in the IKK2-K44M mutant. We had noticed earlier that IKK2-K44M was incapable of de novo autophosphorylation on tyrosine residues (Figure 2D). It also failed to phosphorylate itself on the activation loop Ser177 and 181 (Figure 3A). Interestingly, IKK2-K44M phosphorylated full-length WT IκBα robustly albeit much less efficiency than WT IKK2 (Figure 3B) in a radioactive in vitro kinase assay. IKK2-K44M also showed a minor auto-phosphorylation activity compared to WT IKK2 in the absence of substrate (Figure 3B). We performed a similar kinase assay using full-length WT IκBα where cold ATP was used as the phosphate donor, and phosphorylation status of IκBα at S32, S36 was monitored by immunoblotting using an antibody that specifically recognizes IκBα when phosphorylated at those two N-terminal serines. It revealed that IKK2-K44M is unable to phosphorylate S32/S36 of IκBα (Figure 3C). Since radioactive kinase assays report pan-phosphorylation of a substrate (Figure 3B), it possibly reflects phosphorylation of the PEST region of FL IκBα. We additionally compared IKK2-K44M with WT IKK2 using radioactive kinase assay with both FL IκBα WT and FL IκBαAA substrates in the absence or presence of Inhibitor VII. As anticipated, FL IκBα was phosphorylated much more efficiently than FL IκBαAA by WT IKK2, and phosphorylation was practically abolished in the presence of 10 μM Inhibitor VII (Figure 3D). In contrast, IKK2-K44M phosphorylated FL IκBαAA more efficiently than FL IκBα, and reduction in the presence of inhibitor VII was strikingly less (Figure 3D). It indicates that IKK2-K44M, despite retaining non-specific kinase activity towards the PEST region, lost phosphorylation activity towards S32, S36 of the bona fide substrate IκBα. We subjected 100ug of FL IKK2 K44M protein used in the above studies to LC-MS/MS analysis, which did not detect any contaminating kinase from Sf9 cells. It is noteworthy that an equivalent K to M mutation in Erk2 also did not abolish the basal activity, but rather retained ∼5 % of it(Robbins et al., 1993). The above set of experiments confirm that K44 residue is not only crucial for both IKK2’s well-established serine phosphorylation as well as newly-found tyrosine autophosphorylation activities, as well as its specificity for phosphorylating its bona fide substrate IκBα at S32/36. Hence, these results strongly suggest that IKK2-autophosphorylation is critical for its substrate specificity.

Removed in the revised version: Autophosphorylation of IKK2 is critical to the specificity of IκBα phosphorylation.

Removed in the revised version: (A) Immunoblotting performed with phospho-IKK2 Ser(177/181) antibody to check autophosphorylation of IKK2 mutant K44M at different time points compared to that of IKK2 WT. (B) Comparison of substrate and autophosphorylation activities of FL IKK2 WT and K44M mutant in an in vitro radioactive kinase assay performed with or without FL IκBα WT as the substrate. (C) Residue level specificity of phosphorylation by the FL IKK2 K44M was analysed using an antibody that specifically recognizes phospho-S32/36 of IκBα. (D) Differential sensitivity of specific vs. non-specific phosphorylation by IKK2 as assessed by the effect of Inhibitor VII on FL IKK2 WT and IKK2 K44M using IκBα WT and AA as substrates.

Removed in the revised version: Tyrosine phosphorylation is critically important for IKK2 activity

Removed in the revised version: A number of earlier studies indicated that phosphorylation of IKK2 on Y169, Y188 and/or Y199 to be critical for its function(Darwech et al., 2010; Meyer et al., 2013). Our mass spectrometric analyses identified multiple tyrosine residues to be phosphorylated, e.g., Y169, Y188, Y199, Y325 (data not shown). We compared available structural models of human IKK2 (PDB ID: 4KIK, 4E3C) to investigate whether tyrosine phosphorylation within the activation loop could be supported based on its structural features(Liu et al., 2013; Polley et al., 2013). The structure of IKK2 dimer is asymmetric in that the activation loop of one monomer is largely disordered while that of the other monomer is mostly ordered (PDB ID: 4KIK)(Liu et al., 2013). In the active monomers S177 and S181 are clearly observed. In these active state conformers, the position and conformation of Tyr169 appears to be well-poised for accepting a phosphate from ATP (Figure S4A) intramolecularly. A superposition of IKK2 with a pseudo-substrate-bound PKA shows that the hydroxyl of Tyr169 in IKK2 projects toward the γ-phosphate of ATP similarly to the serine hydroxyl on the PKA inhibitor pseudo-substrate(Knighton et al., 1991; Nolen et al., 2004) (Figure S4A & B). Interestingly, Y169 is the DFG+1 (DLG in case of IKK2 and IKK1) residue while Y188 (equivalent to Y204 of PKA) is located at the end of activation loop prior to the conserved APE motif (Figure 4A and 4B). It was previously established that the DFG+1 position is crucial for defining the substrate specificity of a kinase(Chen et al., 2014). Y188 counterpart Y204 of PKA is reported to allosterically modulate PKA activity by altering its conformational entropy(Ahuja et al., 2017; Yang et al., 2005). We further focused our attention to Y169 and Y188 as they are located in the activation loop of the kinase.

Earlier version: Tyrosine autophosphorylation is critical for IKK2 activity.

Revised version: Dual-specific autophosphorylation is critical for IKK2’s function

Added in the revised version: Previous reports suggested that phosphorylation of IKK2 on Y169, Y188 and/or Y199 was critical for its function(Darwech et al., 2010; Meyer et al., 2013). Our mass spectrometric analyses identified multiple tyrosine residues to be phosphorylated, e.g., Y169, Y188, Y325 (Figure S3A). To glean into the structural plausibility of tyrosine phosphorylation within the activation loop, we compared available structural models of human IKK2 (PDB ID: 4KIK, 4E3C) (Liu et al., 2013; Polley et al., 2013). In the active state conformers, the position and conformation of Tyr169 appears to be well-poised for accepting a phosphate from ATP (Figure S3B) intramolecularly. A superposition of IKK2 with a pseudo-substrate-bound PKA shows that the hydroxyl of Tyr169 in IKK2 projects toward the γ-phosphate of ATP similarly to the serine hydroxyl on the PKA inhibitor pseudo-substrate(Knighton et al., 1991; Nolen et al., 2004) (Figure S3B & C). Interestingly, Y169 is the DFG+1 (DLG in case of IKK2 and IKK1) residue while Y188 (equivalent to Y204 of PKA) is located at the end of activation loop prior to the conserved APE motif (Figure 3A & B). It was previously established that the DFG+1 position is crucial for defining the substrate specificity of a kinase(Chen et al., 2014). Y188 counterpart Y204 of PKA is reported to allosterically modulate PKA activity by altering its conformational entropy(Ahuja et al., 2017; Yang et al., 2005). In order to understand the functional consequence of dual-specific autophosphorylation of IKK2, we further focused our attention to Y169 and Y188 as they are located in the activation loop of the kinase along with the previously described K44M mutant.

Added in the revised version: Dual-specific autophosphorylation is critical for IKK2’s function:

Earlier version: (A) Position of canonically important residues in the AL (purple ribbon) of IKK2-KD are shown in the context of its structure (PDB ID 4E3C; surface representation; KD is shown light green with respect to the SDD in teal). Autophosphorylated tyrosine residues identified in this study are shown in green. (B) Alignment of the activation segment sequences of different IKK family kinases. Conservation of Tyr at the DFG+1 (DLG in case of IKK1 and IKK2) is retained only in a stress response related plant kinase SnRK2, but not in other IKK homologues or any known dual specificity kinases, e.g., DYRK and GSK3β (both contain Ser at that position). Tyr188 (204 in PKA) is conserved. (C) Radioactive in vitro kinase assay performed with IKK2 Y169F and Y188F mutants using GST (1-54) IκBα WT and AA as substrates. (D) Radioactive in vitro kinase assay performed with IKK2 Y169F mutant with a range of GST (1-54) IκBα mutants where the substrate phosphorylation signal is abolished in case of S36 A/E IκBα. (E) Severe reduction in TNFα-induced IKK activity in ikk2-/- MEF-3T3 cells reconstituted with mutant (Y169F and Y188F) IKK2s compared to wild-type. Kinase assay was performed with IKK immunoprecipitated (IP-ed) with anti-NEMO antibody from whole cell extract. (F) AL-serine phosphorylation status of IKK2 WT and Y169F was assessed in reconstituted MEFs post-stimulation with TNF-α after immunoprecipitating the kinase with monoclonal anti-HA antibody.

Revised version: (A) Position of canonically important residues in the AL (purple ribbon) of IKK2-KD are shown in the context of its structure (PDB ID 4E3C; surface representation; KD is shown light green with respect to the SDD in teal). Autophosphorylated tyrosine residues identified in this study are shown in green. (B) Alignment of the activation segment sequences of different IKK family kinases. Conservation of Tyr at the DFG+1 (DLG in case of IKK1 and IKK2) is retained only in a stress response related plant kinase SnRK2, but not in other IKK homologues or any known dual specificity kinases, e.g., DYRK and GSK3β (both contain Ser at that position). Tyr188 (204 in PKA) is conserved. (C) Comparison of substrate and autophosphorylation activities of FL IKK2 WT and K44M mutant in an in vitro radioactive kinase assay performed with or without FL IκBα WT as the substrate. (D) Residue level specificity of phosphorylation by the FL IKK2 K44M was analysed using an antibody that specifically recognizes phospho-S32/36 of IκBα. (E) Radioactive in vitro kinase assay performed with IKK2 Y169F and Y188F mutants using GST (1-54) IκBα WT and AA as substrates. (F) Radioactive in vitro kinase assay performed with IKK2 Y169F mutant with a range of GST (1-54) IκBα mutants where the substrate phosphorylation signal is abolished in case of S36 A/E IκBα. (G) Severe reduction in TNFα-induced IKK activity in ikk2-/- MEF-3T3 cells reconstituted with mutant (Y169F and Y188F) IKK2s compared to wild-type. Kinase assay was performed with IKK immunoprecipitated (IP-ed) with anti-NEMO antibody from whole cell extract (n=2).

Added in the revised version: Firstly, we wanted to check the effect of autophosphorylation on IKK2 activity by using the kinase inactivating K44M mutant that was found to be incapable of tyrosine autophosphorylation (Figure 2D) and fails to undergo autophosphorylation of the activation loop serines (Figure S3D). Interestingly, IKK2-K44M showed a very weak auto-phosphorylation activity compared to WT IKK2 but phosphorylated full-length WT IκBα robustly albeit much less efficiently than WT IKK2 in a radioactive in vitro kinase assay (Figure 3C). However, IKK2-K44M was unable to specifically phosphorylate S32/S36 of IκBα (Figure 3D). These seemingly contrasting results possibly reflects phosphorylation of the PEST region of FL IκBα in the radioactive kinase assays as they report pan-phosphorylation of a substrate (Figure 3C & D). In addition, IKK2-K44M phosphorylated FL IκBαAA more efficiently than FL IκBα WT unlike WT IKK2 that phosphorylated FL IκBα WT much more efficiently than FL IκBαAA (Figure S3E). It indicates that IKK2-K44M retained non-specific kinase activity to some extent while lost functionally relevant specificity towards the S32, S36 of IκBα. We subjected 20ug of FL IKK2 K44M protein used in the above studies to LC-MS/MS analysis, which did not detect any significant contaminating kinase from Sf9 cells. It is noteworthy that an equivalent K to M mutation in Erk2 also did not abolish the basal activity, but rather retained ∼5 % of it(Robbins et al., 1993).

Earlier version: To test the effect of phosphorylation at Y169 and Y188, we generated Y169F and Y188F mutants of IKK2. Radioactive in vitro kinase assays using recombinant purified IKK2FL-Y188F and IKK2FL-Y169F mutants revealed a drastic reduction in the activity of IKK2 (Figure 4C). Since the DFG+1 residue in kinases (Y169 in IKK2) is known to dictate its specificity, we created different mutants of the IκBα substrate where S32 and S36 were singly and independently altered to either alanine (phosphoablative) or glutamic acid (phosphomimetic) residues. Kinase activity of IKK2-Y169F was monitored in an in vitro radioactive kinase assay using WT as well as mutant IκBα. Mutation of S36 to either alanine or glutamate drastically reduced the phosphorylation signal of IκBα while that for the S32A/E mutants were minimally affected (Figure 4D). This suggests that the presence of Tyr (and possibly its phosphorylation) is critically important for signal-responsive phosphorylation of IκBα.

Revised version: Next, we wanted to understand the functional relevance of tyrosine autophosphorylation using two tyrosine-phosphoablative mutants of IKK2 (FL IKK2-Y188F and FL IKK2-Y169F), both of which showed severe reduction in activity in an in vitro kinase assay (Figure 3E). Since the DFG+1 residue in kinases (Y169 in IKK2) is known to dictate its specificity, we created different single mutants of IκBα where S32 and S36 were independently altered to either alanine (phosphoablative) or glutamic acid (phosphomimetic) residues. FL IKK2-Y169F displayed drastically reduced level of kinase activity towards S36A/E mutants of IκBα while that for the S32A/E mutants were minimally affected compared to the WT substrate (Figure 3F). Reduction of phosphorylation S36A/E by the WT IKK2 was less severe compared to that by the Y169F mutant kinase (Figure S3F). The above observations suggest that the presence of Tyr (and possibly its phosphorylation) is critically important for signal-responsive phosphorylation of IκBα.

Earlier version: Next, we tested the consequence of Y169F and Y188F mutants on the pro-inflammatory signal-induced IKK activity in cells. We reconstituted IKK2 knock-out mouse embryonic fibroblast (ikk2-/- MEF) cells with WT, Y169F and Y188F mutants of IKK2, and measured signal-induced activation of IKK2 upon TNF-α treatment. IKK2 activity is significantly reduced in Y169F and nearly abolished in Y188F (Figure 4E). Interestingly, phosphorylation of activation loop serines of the less active Y169F mutant was only marginally defective (Fig. 4F). This suggests that phosphorylation of the activation loop serines at 177 and 181 is necessary but not sufficient for IKK2 to become fully active. Together, these results suggest an unreported communication between the phosphorylation of Y169 of IKK2 and that of S32 of IκBα.

Revised version: Next, we reconstituted IKK2 knock-out mouse embryonic fibroblast (ikk2-/- MEF) cells with WT, Y169F and Y188F mutants of IKK2, and measured signal-induced activation of IKK2 upon TNF-α treatment. IKK2 activity is significantly reduced in Y169F and nearly abolished in Y188F (Figure 3G). Phosphorylation of activation loop serines of the less active Y169F mutant, nonetheless, was only marginally defective (Figure S3G) suggesting that activation loop phosphorylation is necessary but not sufficient for IKK2 to become fully active. Together, these results hint at an unreported communication between the phosphorylation of Y169 of IKK2 and that of S32 of IκBα.

Structural analyses of IKK2 autophosphorylation

Earlier version: To get a more comprehensive picture, it is necessary to analyze at least three phosphorylated states of IKK2: unphosphorylated (UnP-IKK2), phosphorylated at S177 and S181 (p-IKK2), and phosphorylated at Y169, S177 and S181 (P-IKK2). We took advantage of computational approaches using molecular dynamic (MD) simulations and flexible molecular docking (see methods).43–50 MD simulation was performed for 200ns for each of these three states of IKK2, unphosphorylated and phosphorylated. Phosphorylation at S177, S181 increased the stability of the kinase, which was further increased by the phosphorylation at Y169 as evidenced by a gradual decrease in the total energy of system (Figure 5A). A striking enhancement in stability of IKK2 upon this auto-phosphorylation is reflected in the change of Tm from ∼ 40 °C to ∼ 50 °C in differential scanning calorimetry (DSC) scans in the presence of ADP and ATP (Figure S4A). This behavior supports our observation in MD simulation, and tallies with increased solubility of IKK2 observed after auto-phosphorylation that was essential for its successful crystallization (Polley et al., 2013)

Revised version: To get a more comprehensive picture, it is necessary to analyze at least three phosphorylated states of IKK2: unphosphorylated (UnP-IKK2), phosphorylated at S177 and S181 (p-IKK2), and phosphorylated at Y169, S177 and S181 (P-IKK2). We took advantage of computational approaches using molecular dynamic (MD) simulations and flexible molecular docking (see methods). MD simulation was performed for 200ns for each of these three states of IKK2, unphosphorylated and phosphorylated. Phosphorylation at S177, S181 increased the stability of the kinase, which was further increased by the phosphorylation at Y169 as evidenced by a gradual decrease in the total energy of system (Figure 4A). A striking enhancement in stability of IKK2 upon this auto-phosphorylation is reflected in the change of Tm from ∼ 40 °C to ∼ 50 °C in differential scanning calorimetry (DSC) scans in the presence of ADP and ATP (Figure S4A). This behavior supports our observation in MD simulation, and tallies with increased solubility of IKK2 observed after auto-phosphorylation that was essential for its successful crystallization (Polley et al., 2013).

Structural analyses of IKK2 autophosphorylation.

(A) Molecular dynamics (MD) simulations of all three differently phosphorylated states of IKK2 was performed. 200ns trajectory of energy for each state is shown in golden yellow, blue and copper red representing UnP-IKK2, p-IKK2 and P-IKK2, respectively. Same color coding has been followed throughout this figure and the corresponding supplementary figure. (B) 200ns trajectory of RMSD for each state is shown. (C) Residues forming canonical R-(labelled in blue) & C-spines (labelled in dark gray) representative of an active conformation in p-IKK2 are shown. (D) Superposition of structures representing the differently phosphorylated states post 200ns MD simulation. Relevant sections, e.g., activation loop, αC-helix, Gly-rich loop areas are highlighted in the right panel. (E) ATP-bound p-IKK2 structure is depicted. In the right panel, relevant residues are highlighted. Y169 is found in a position poised for autophosphorylation. (F) ATP completes the desired continuum of R-& C-spine observed in active kinase conformations. (G) P-IKK2 cannot accommodate ATP in its cleft, and αC-helix is displaced. (H) M is moved away from the R-spine, failing to form canonically active conformation of R-& C-spines. (I & J) ADP-bound p-IKK2 and P-IKK2 states are shown, respectively. Both p-IKK2 and P-IKK2 states can accommodate ADP in their respective clefts.

Earlier version: Interestingly, RMSD values increased with more phosphorylation (Figure 5B) despite phosphorylation-induced stabilization, P-IKK2 showed the highest RMSD followed by p-IKK2 and UnP-IKK2, indicating increased internal motion in the kinase as a result of phosphorylation. It may be noted that, the changes in total energy and RMSD are consistent throughout the course of simulation though the magnitude is small. We found that, phosphorylation at S177, S181 and S177, S181, Y169 induced distinct structural alteration in the KD when compared with the unphosphorylated version (Figure 5C, models shown separately in Figure S4B). Careful analyses of the trajectory of RMSF values for these three states revealed that, different regions of the kinase domain had different RMSF values in different phosphorylated states (Figure S4C). We observed that, while the glycine-rich loop and the αC-helix in p-IKK2 were situated in a manner reminiscent of an active kinase (Figure S4D), those in UnP and P-IKK2 were placed very differently. In fact, helix αC was found to be distorted in P-IKK2, that is unlikely to support canonical phosphotransfer from ATP. Activation segment/loop (activation loop) in these three differently phosphorylated states adopted conformations distinct from each other. activation loop in the UnP model was displayed out and moved closer to the tip of the C-lobe, whereas activation loop in p-IKK2 and P-IKK2 moved inwards and adopted conformations closer to the N-lobe. The formation of the proper R- and C-spines in p-IKK2 confirmed its active form (Figure 5D). p-IKK2 KD also exhibited other canonical features of an active kinase, viz., a salt bridge between K44 and E61 (K72, E91 in PKA, respectively), ‘DFG in’ (DLG here) conformation. Next, we docked ATP onto the p-IKK2 structure (ΔG = -10.64 kcal/mol) using a flexible docking protocol (see methods). The ATP and its phosphate groups exhibited a pose (Figure S4D) very similar to that in the PKA structure (PDB ID: 1ATP), and the terminal phosphate was in close proximity to Y169-OH making autophosphorylation of Y169 highly plausible (Figure 5E). Further analyses of the ATP-bound p-IKK2 structure confirm that presence of ATP helped the R and C-spine form the desired continuum seen in the active form of a kinase (Figure 5F). Superposition of the ATP from the p-IKK2 on P-IKK2 indicated that the narrowing of ATP-binding pocket in P-IKK2 may lead to rejection of ATP emanating from severe clashes between the glycine-rich loop and ATP (Figure 5G). Furthermore, the R-spine in P-IKK2 was not continuous (Figure 5H), wherein M65 (L95 in PKA) of the four conserved residues were located at a position far away from the rest of three residues that constitute the R-spine, and from the glycine rich loop. In addition, dynamic cross-correlation matrix (DCCM) or contact map of each structure suggests that specific phosphorylation events render the kinase display distinct allosteric changes at locations far from the phosphorylation sites themselves (Figure S4E).

Revised version: Interestingly, RMSD values increased with more phosphorylation (Figure 4B) despite phosphorylation-induced stabilization, P-IKK2 showed the highest RMSD followed by p-IKK2 and UnP-IKK2, indicating increased internal motion in the kinase as a result of phosphorylation. It may be noted that, the changes in total energy and RMSD are consistent throughout the course of simulation though the magnitude is small. We found that, phosphorylation at S177, S181 and S177, S181, Y169 induced distinct structural alteration in the KD when compared with the unphosphorylated version (Figure 4C, models shown separately in Figure S4B). Careful analyses of the trajectory of RMSF values for these three states revealed that, different regions of the kinase domain had different RMSF values in different phosphorylated states (Figure S4C). We observed that, while the glycine-rich loop and the αC-helix in p-IKK2 were situated in a manner reminiscent of an active kinase, those in UnP and P-IKK2 were placed very differently (Figure 4C). In fact, helix αC was found to be distorted in P-IKK2, that is unlikely to support canonical phosphotransfer from ATP. Activation segment/loop (activation loop) in these three differently phosphorylated states adopted conformations distinct from each other: activation loop in the UnP model was splayed out and moved closer to the tip of the C-lobe, whereas activation loop in p-IKK2 and P-IKK2 moved inwards and adopted conformations closer to the N-lobe (Figure 4C). The formation of the proper R-and C-spines in p-IKK2 confirmed its active form (Figure 4D). p-IKK2 KD also exhibited other canonical features of an active kinase, viz., a salt bridge between K44 and E61 (K72, E91 in PKA, respectively), ‘DFG in’ (DLG here) conformation. In addition, dynamic cross-correlation matrix (DCCM) or contact map of each structure suggests that specific phosphorylation events render the kinase display distinct allosteric changes at locations far from the phosphorylation sites themselves (Figure S4D).

Added in the revised version: Next, we docked ATP onto the these three phosphorylated-state structures using LeDock and GOLD followed by rescoring with AutoDock Vina using the 0 ns (starting, S) and 200ns (end, E) structures from MD simulation (see methods). It was observed that ATP binding to P-IKK2 is relatively more unfavorable (docking score in positive range) both in the starting and end structures (Figure S4E), whereas UnP and p-IKK2 binds to ATP favorably. Interestingly, each of the phospho-/unphosphorylated IKK2 species binds to ADP favorably except for the P-IKK2 in the starting conformation (Figure S4E). To look into these complexes more carefully, we extracted 50 intermediate structures from a 10 ns MD simulation and calculated binding free energies (ΔG) of these intermediate structures using the MM-PBSA (Molecular Mechanics Poisson-Boltzmann Surface Area) method. Clearly, ATP-binding is highly unfavourable for the P-IKK2 population that has high relative preference for ADP, while p-IKK2 shows comparable preference for both (Figure S4F). Among the docked structures, ATP and its phosphate groups exhibited a pose (ΔG = -10.64 kcal/mol) (Figure S4G) very similar to that in the PKA structure (PDB ID: 1ATP), and the terminal phosphate was in close proximity to Y169-OH making autophosphorylation of Y169 highly plausible (Figure 4E). Further analyses of the ATP-bound p-IKK2 structure confirm that presence of ATP helped the R and C-spine form the desired continuum seen in the active form of a kinase (Figure 4F). Superposition of the ATP from the p-IKK2 on P-IKK2 indicated that the narrowing of canonical ATP-binding pocket in P-IKK2 may lead to rejection of ATP from the canonical binding pocket emanating from severe clashes between the glycine-rich loop and ATP (Figure 4G). However, this does not exclude the possibility of P-IKK2 employing alternative binding modes to interact with ATP. Furthermore, the R-spine in P-IKK2 was not continuous (Figure 4H), wherein M65 (L95 in PKA) of the four conserved residues were located at a position far away from the rest of three residues that constitute the R-spine, and from the glycine rich loop.

Earlier version: ATP binding to p-IKK2 leads to autophosphorylation at Y169 generating P-IKK2 in an ADP bound state. Interestingly, the superposition of ADP (from the docked complex of p-IKK2 and ADP, ΔG = -9.88 kCal/mol, Figure 5I) onto the P-IKK2 structure confirmed that while P-IKK2 was unable to accommodate ATP in a manner similar to that in other known experimentally derived kinase structures, it could comfortably accommodate ADP into the ATP-binding cleft (Figure 5J). It may be noted that, p-IKK2 has higher affinity for ATP than ADP as indicated by their respective binding energy values. We speculate that ADP-bound P-IKK2 is an intermediate step where presence of as many phosphate groups in the activation loop and ADP in such a close proximity persuades the kinase to adopt a conformation that can withstand this high density of negative charges albeit risking the integrity of an active fold, and making some of those de novo phosphate groups transient in nature to return back to the p-IKK2 state to readily bind to ATP that is abound in the reaction.

Revised version: ATP binding to p-IKK2 leads to autophosphorylation at Y169 generating P-IKK2 in an ADP bound state. Interestingly, the superposition of ADP (from the docked complex of p-IKK2 and ADP, ΔG = -9.88 kCal/mol, Figure 4I) onto the P-IKK2 structure confirmed that while P-IKK2 was unable to accommodate ATP in a manner similar to that in other known experimentally derived kinase structures, it could comfortably accommodate ADP into the ATP-binding cleft (Figure 4J). It may be noted that, p-IKK2 has comparable affinities for ATP and ADP, like other kinases (Becher et al., 2013), as indicated by their respective binding energy/docking score values. Worth noting, cellular concentration of ATP and ADP plays important role in the outcome of this process where ATP concentration is much higher than ADP. We speculate that ADP-bound P-IKK2 is an intermediate step where presence of as many phosphate groups in the activation loop and ADP in such a close proximity persuades the kinase to adopt a conformation that can withstand this high density of negative charges albeit risking the integrity of an active fold, and making some of those de novo phosphate groups transient in nature to return back to the p-IKK2 state to readily bind to ATP that is abound in the reaction.

These structural observations suggest that it is possible for the residues within the flexible activation segment to undergo auto-phosphorylation and possibly render the kinase with a novel phospho-transfer activity distinct from the conventional transfer of γ-phosphate from ATP to the substrate.

Freshly auto-phosphorylated IKK2 phosphorylates IκBα even in absence of ATP

Earlier version: We observed phosphorylation of tyrosine residue(s) only upon fresh ATP-treatment, indicating that tyrosines of purified IKK2 loses its phosphate group during the course of purification (Eckhart et al., 1979). It may be noted that, we did not observe any density of phosphates on any of these residues in the human IKK2 crystal structure derived from ATP-treated protein sample (Polley et al., 2013). Of note, human IKK2 crystals grown with autophosphorylated IKK2 (PDB ID 4E3C) took about a month at 18-20°C before they were subjected to X-ray diffraction. This suggested an intriguing possibility that the phosphate linked to Tyr residue(s) may serve as a temporary phosphate-sink and could eventually be transferred to substrate IκBα. To test this possibility, we first allowed purified IKK2 to undergo auto-phosphorylation by incubating the excess kinase with limiting amounts of radiolabeled γ-32P-ATP followed by two passes through desalting spin columns (40 kDa MW cut-off) to remove unreacted ATP. Nucleotide-free auto-phosphorylated IKK2 (P-IKK2) was incubated with IκBα either in absence of nucleotide or with newly added unlabeled nucleotides, and phosphorylation status of the substrate was subsequently monitored by autoradiography after resolving the reaction mixture on an SDS-PAGE (Figure 6A). Transfer of radiolabeled phosphate to IκBα was observed and the transfer was significantly enhanced in the presence of cold ADP (Figure 6B). Phosphorylation of IκBα did not significantly decrease in the presence of a high concentration of cold ATP. That addition of excess cold ATP did not reduce the extent of IκBα phosphorylation-signal on the autoradiograph strongly suggests the transfer of 32P-phosphate from 32P-IKK2 to IκBα (Figure 6B) by a relay mechanism. This transfer perhaps occurs through a phosphoenzyme intermediate with or without involving ADP.

Revised version: We observed phosphorylation of tyrosine residue(s) only upon fresh ATP-treatment, indicating that tyrosines of purified IKK2 loses its phosphate group during the course of purification (Eckhart et al., 1979). It may be noted that, we did not observe any density of phosphates on any of these residues in the human IKK2 crystal structure derived from ATP-treated protein sample (Polley et al., 2013). Of note, human IKK2 crystals grown with autophosphorylated IKK2 (PDB ID 4E3C) took about a month at 18-20°C before they were subjected to X-ray diffraction. This suggested an intriguing possibility that the phosphate linked to Tyr residue(s) may serve as a temporary phosphate-sink and could eventually be transferred to substrate IκBα. To test this possibility, we first allowed purified IKK2 to undergo auto-phosphorylation by incubating the excess kinase with limiting amounts of radiolabeled γ-32P-ATP followed by two passes through desalting spin columns (40 kDa MW cut-off) to remove unreacted ATP. Nucleotide-free auto-phosphorylated IKK2 (P-IKK2) was incubated with IκBα either in absence of nucleotide or with newly added unlabeled nucleotides, and phosphorylation status of the substrate was subsequently monitored by autoradiography after resolving the reaction mixture on an SDS-PAGE (Figure 5A). Of note, such an experimental scheme essentially captures a single turnover phosphorylation of the substrate IκBα. Transfer of radiolabeled phosphate to IκBα was observed and the transfer was significantly enhanced in the presence of cold ADP (Figure 5B). Phosphorylation of IκBα did not significantly decrease in the presence of a high concentration of cold ATP. That addition of excess cold ATP did not reduce the extent of IκBα phosphorylation-signal on the autoradiograph strongly suggests the transfer of 32P-phosphate from 32P-IKK2 to IκBα (Figure 5B) by a relay mechanism. This transfer perhaps occurs through a phosphoenzyme intermediate with or without involving ADP.

Freshly autophosphorylated IKK2 relays phosphates to IκBα.

(A) Schematic representation of the experiment performed to monitor the path traversed by phosphate(s) from phospho-IKK2 to substrate by autoradiography detection method. (B) Purified autophosphorylated (radiolabelled) IKK2 transfers its phosphate to IκBα substrate in absence of any nucleotide, and transfer efficiency is enhanced upon addition of ADP or ATP. (C) Schematic representation of the experiment performed to monitor the path traversed by phosphate(s) from phospho-IKK2 to substrate by immunoblotting detection method. (D) Elution profile of the size exclusion chromatography run of phospho-IKK2 to get rid of excess unlabelled ATP performed in Superdex200 10/30 increase. (E) Purified autophosphorylated (cold) IKK2 transfers its phosphate to IκBα substrate with enhanced transfer efficiency upon addition of ADP or ATP monitored by immunoblotting using specific antibodies as indicated in the figure.

Earlier version: We undertook a different strategy to further confirm transfer of phosphate from the kinase (and not ATP) to the substrate. We auto-phosphorylated IKK2FL-WT with excess cold ATP and removed the unreacted ATP from the reaction mixture using size-exclusion chromatography on a Superdex200 column (Figure 6C). P-IKK2 eluted at fractions corresponding to MW between 670 and 158 kDa, i.e., much larger than free ATP (Figure 6D). This P-IKK2 protein was incubated with FL IκBα either in presence or absence of ADP for different time periods. A similar reaction in presence of ATP instead of ADP was performed as a positive control. Phosphorylation of S32 and S36 (detected by the phospho-serine antibody specific to S32 and S36 of IκBα) was efficient in presence of 10 and 50 μM ADP but absent when ADP was absent in the reaction mixture. Phosphorylation of IκBα in the presence of 50 μM ATP was significantly more robust as anticipated (Figure 6E). The minor detection of IκBα phosphorylation in absence of ADP in the radioactive assay and lack of it in the cold immunoblotting assay (compare lanes marked with vertical arrows in Figures 6B and 6F) may indicate higher sensitivity of the radioactive assay. The need of ADP in phosphotransfer prompted us to consider whether the ADP was contaminated with a trace amount of ATP. ESI-MS analysis of 50 μM ADP did not detect any peak corresponding to ATP (Figure S5A). Dependence of IκBα phosphorylation on ADP could raise another possibility if microscopic reversibility generated ATP from ADP in the active site of the kinase. To this end, we employed TLC-based analysis to detect generation/presence of ATP in a similar setting as described in Fig. 5A-B which revealed that no detectable ATP was present in the system, i.e., no detectable ATP was retained in the protein after passing it through the spin column twice, nor did any γP32-ATP was generated when incubated with cold ADP (Figure S5B).

Revised version: We undertook a different strategy to further confirm transfer of phosphate from the kinase (and not ATP) to the substrate. We auto-phosphorylated IKK2FL-WT with excess cold ATP and removed the unreacted ATP from the reaction mixture using size-exclusion chromatography on a Superdex200 column (Figure 5C). P-IKK2 eluted at fractions corresponding to MW between 670 and 158 kDa, i.e., much larger than free ATP (Figure 5D). This P-IKK2 protein was incubated with FL IκBα either in presence or absence of ADP for different time periods. A similar reaction in presence of ATP instead of ADP was performed as a positive control. Phosphorylation of S32 and S36 (detected by the phospho-serine antibody specific to S32 and S36 of IκBα) was efficient in presence of 10 and 50 μM ADP but absent when ADP was absent in the reaction mixture. Phosphorylation of IκBα in the presence of 50 μM ATP was significantly more robust as anticipated (Figure 5E). The minor detection of IκBα phosphorylation in absence of ADP in the radioactive assay and lack of it in the cold immunoblotting assay (compare lanes marked with vertical arrows in Figures 5B and 5F) may indicate higher sensitivity of the radioactive assay. The need of ADP in phosphotransfer prompted us to consider whether the ADP was contaminated with a trace amount of ATP. ESI-MS analysis of 50 μM ADP did not detect any peak corresponding to ATP (Figure S5A). Dependence of IκBα phosphorylation on ADP could raise another possibility if microscopic reversibility generated ATP from ADP in the active site of the kinase. To this end, we employed TLC-based analysis to detect generation/presence of ATP in a similar setting as described in Figure 5A & B which revealed that no detectable ATP was present in the system, i.e., no detectable ATP was retained in the protein after passing it through the spin column twice, nor did any γP32-ATP was generated when incubated with cold ADP (Figure S5B).

Phosphotransfer from P-IKK2 is critical to fidelity and specificity of IκBα phosphorylation

Earlier version: We next tested if phosphotransfer is restricted only to the N-terminal signal responsive S32/S36 or to serines/threonines of the PEST domain as well. Radioactive kinase assay as described earlier was performed with full-length WT and mutant S32A, S36A (AA) IκBα as the phosphate recipient. Phosphate group(s) of P-IKK2 was not transferred to full-length IκBα AA which retains the serines and threonines within its PEST domain (Figure 7A), indicating phosphotransfer only to the N-terminal serines of IκBα that are S32 and S36. We also tested the effect of AMPPNP, an unhydrolyzable analogue of ATP, on the phosphotransfer, and found no inhibitory effect irrespective of the presence of ADP. This suggests adenine nucleotides poising the kinase for phosphotransfer (Figure 7A). A recent report proposed that IKK2 phosphorylates IκBα at S32 followed by S36 following a single binding event and that phosphorylation of S32 increases the phosphorylation rate of S36(Stephenson et al., 2022). Based on our results and other evidence, we propose that IKK2 phosphorylated at S177/S181 undergoes autophosphorylation on tyrosine residues in its activation loop area that serves as a phosphoenzyme intermediate to transfer phosphate group from IKK2 to S32 that enhances the rate of second site (S36) phosphorylation (Figure 7B). For the second site phosphorylation ATP serves as the phosphate donor.

Revised version: We next tested if phosphotransfer is restricted only to the N-terminal signal responsive S32/S36 or to serines/threonines of the PEST domain as well. Radioactive kinase assay as described earlier was performed with full-length WT and mutant S32A, S36A (AA) IκBα as the phosphate recipient. Phosphate group(s) of P-IKK2 was not transferred to full-length IκBα AA which retains the serines and threonines within its PEST domain (Figure 6A), indicating phosphotransfer only to the N-terminal serines of IκBα that are S32 and S36. We also tested the effect of AMPPNP, an unhydrolyzable analogue of ATP, on the phosphotransfer, and found no inhibitory effect irrespective of the presence of ADP. This suggests adenine nucleotides poising the kinase for phosphotransfer (Figure 6A). A recent report proposed that IKK2 phosphorylates IκBα at S32 followed by S36 following a single binding event and that phosphorylation of S32 increases the phosphorylation rate of S36(Stephenson et al., 2022). Based on our results and other evidence, we propose that IKK2 phosphorylated at S177/S181 undergoes autophosphorylation on tyrosine residues in its activation loop area that serves as a phosphoenzyme intermediate to transfer phosphate group from IKK2 to S32 that enhances the rate of second site (S36) phosphorylation (Figure 6B). For the second site phosphorylation ATP serves as the phosphate donor.

Specificity and fidelity of phosphotransfer by IKK2 to IκBα.

Earlier version: (A) AMP-PNP can support efficient phosphotransfer. Also, phosphotransfer is observed in IκBα WT but not in S32,36Ala double mutant. Domain organization and position of relevant S/T/Y residues of IκBα are shown above the autoradiograph. (B) A detailed scheme of reactions showing IKK2-activation and subsequent phosphorylation of the signal responsive S32,36 on IκBα by phosphotransfer in presence of ADP. (C) A cartoon representation of the proposed model of NF-κB activation by IKK2 employing phosphotransfer to IκBα.

Revised version: (A) AMP-PNP can support efficient phosphotransfer. Also, phosphotransfer is observed in IκBα WT but not in S32,36Ala double mutant. Domain organization and position of relevant S/T/Y residues of IκBα are shown above the autoradiograph. (B) A detailed scheme of reactions showing IKK2-activation and subsequent phosphorylation of the signal responsive S32,36 on IκBα by phosphotransfer in presence of ADP.

Discussion

Earlier version: Phosphorylation of substrates by protein kinases is regulated in various ways(Cullati et al., 2022; Sang et al., 2022). IKK2, a kinase central to inflammation, phosphorylates serine residues of its primary substrate IκBα with γ-phosphate of ATP. Here, we report multiple intriguing properties of IKK2 that provide specificity to its substrate phosphorylation. These include: 1) dual specificity that entails autophosphorylation of tyrosine residues of IKK2, in addition to known auto-phosphorylation at activation loop serines; 2) relay of phosphate to substrate IκBα from auto-phosphorylated IKK2 (P-IKK2) and 3) loss of specificity of the kinase due to disruption of a universally conserved salt bridge (mediated by K44).

Revised version: We report several intriguing properties of IKK2 that contributes to its substrate phosphorylation specificity and efficiency: 1) dual specificity, involving autophosphorylation of several tyrosine residues of IKK2, in addition to known auto-phosphorylation at activation loop serines; 2) phosphate relay to the substrate IκBα via Tyr169 of auto-phosphorylated IKK2 (P-IKK2); and 3) loss of substrate specificity, despite retained catalytic activity, upon disruption of a universally conserved salt bridge mediated by Lys44.

Added in the revised version: Dual specificity of IKK2

Added in the revised version: Our present analyses of IKK2, however, revealed its surprising property wherein IKK2 phosphorylated at activation loop-serines (S177/S181) could be further auto-phosphorylated by ATP at multiple other residues yielding another form (P-IKK2) (Figure 6B). This marks IKK2 as an autocatalytic dual-specificity kinase rather than a prototype S/T kinase. This multisite phosphorylated form of IKK2 (P-IKK2) could transfer phosphate(s) to substrate IκBα in presence of ADP without requiring exogenous supply of fresh ATP (Figure 5B & E), hinting at an unusual property of IKK2 acting possibly as a phosphate-sink.

Added in the revised version: The phospho relay mechanism

Added in the revised version: Through various experimental approaches, we demonstrated that S32 of IκBα is phosphorylated via a relay mechanism. The transfer of phosphate from Tyr169 to Ser32 of IκBα represents an intriguing and unprecedented mechanism among metazoan protein kinases. We propose a plausible explanation for process: The hydroxyl group of Tyr169 is optimally positioned to readily accept the γ-phosphate from ATP at the expense of minimum energy in presence or absence of the substrate IκBα. The energy required to transfer the phosphate from Tyr169 to Ser32 is lower than that needed for a direct transfer of the same phosphate from ATP to Ser32. Thus, this relay pathway – from ATP to Tyr169 to Ser32 – might offer a more energetically favorable pathway than a direct transfer from ATP to the substrate. Another intriguing feature of phosphorelay in IKK2 is IKK2’s dependence on ADP (Figure 5B & E), which the generic kinase inhibitor AMPPNP (unhydrolyzable ATP analogue) failed to inhibit (Figure 6A). Taken together, we surmise that presence of ADP or AMP-PNP helps IKK2 attain a conformational state poised for phosphorylating S32/36 of IκBα even in absence of fresh exogenous ATP. Similar phenomenon was reported for IRE1, where ADP was found to be a more potent stimulator of its ribonuclease activity and that AMPPNP could also potentiate this activity in IRE1(Lee et al., 2008; Sidrauski and Walter, 1997). We speculate that Ser32 cannot be efficiently phosphorylated by p-IKK2:ATP due to a deliberate chemical mismatch of the hydroxyl group of Ser32 within the active site. However, it aligns well with the P-IKK2:ADP intermediate. To validate this model, further structural studies and MD simulations studies are necessary. Indirect supporting evidence, as described below, underscores the existence of a unique phospho-transfer mechanism.

Added in the revised version: Tyr169 is the only residue surrounding the active site that can be involved in this relay. As phosphorylation of S32, S36 of IκBα is observed with auto-phosphorylated IKK2 EE, transferred phosphates are not from S177 or S181. Y169 is uniquely located at the DFG+1 (DLG in IKK2) position of the magnesium binding loop. This residue is exclusive to IKK1 and IKK2 and is absent even in close mammalian homologues, IKKε and TBK1. Notably, neither IKKε nor TBK1 can phosphorylate IκBα at S32 and S36, either in vitro or in vivo(Larabi et al., 2013; Tu et al., 2013) supporting the assertion that the phospho-transfer is a genuine and specific event. Recent findings suggest that the DFG+1 residue plays a critical in distinguishing Ser versus Thr residues (Chen et al., 2014). These observations highlight the dual role of Y169 in both substrate selectivity and catalytic activity. This dual selectivity – at residue level (serine versus threonine) and the catalytic level (mediated by relay Tyr169) underpins the exquisite specificity of IκBα for the IKK. Additional mechanisms, such as conformational changes associated with Tyr169 phosphorylation may further enhance the efficiency of the relay by properly positioning Ser32 for phospho-transfer. It is worthwhile to note that, a phospho-resistant mutation at S32 in IκBα to Ile (S32I) is associated with ectodermal dysplasia and T-cell immunodeficiency in patients(Courtois et al., 2003; Mooster et al., 2015). To the best of our knowledge, it is the first report, apart from MHCK of Dictyostelium discoideum (an atypical EPK), to suggest that a metazoan EPK confirms its specificity through the formation of a phosphoenzyme intermediate(Ye et al., 2010).

Removed in the revised version: Ser/Thr/Tyr family of eukaryotic protein kinases transfer γ-phosphate of ATP directly to substrate, in general to S, T and Y residues(Cohen, 2002; Fischer and Krebs, 1955; Hunter, 1991; Krebs and Fischer, 1955; Nolen et al., 2004). Histidine kinases (in prokaryotes and in some lower eukaryotes) however shuttles the γ-phosphate to Asp residue on a response regulator (RR) substrate through a phospho-His intermediate(Borkovich and Simon, 1990; Gao and Stock, 2009; Hunter, 2022; Kalagiri and Hunter, 2021; Laub et al., 2007; Robbins et al., 1993). Our present analyses of eukaryotic protein kinase IKK2 revealed its surprising property in which IKK2 phosphorylated at activation loop-serines (S177/S181) could be further auto-phosphorylated by ATP at multiple other residues yielding another form (P-IKK2) (Figure 7B). This marks IKK2 as an autocatalytic dual-specificity kinase rather than a prototype S/T kinase. This multisite phosphorylated form of IKK2 (P-IKK2) could transfer phosphate(s) to substrate IκBα in presence of ADP without requiring exogenous supply of fresh ATP, hinting at an unusual property of IKK2 acting possibly as a phosphate-sink. It is possible that phospho-tyrosines residues of IKK2 act as this sink. The Y169 located at the DFG+1 (DLG in IKK2) position is unique only to IKK1 and IKK2, and is absent even in close mammalian homologues such as IKKε and TBK1. Interestingly, neither IKKε nor TBK1 can phosphorylate IκBα at S32 and S36, in vitro or in vivo(Larabi et al., 2013; Tu et al., 2013). A recent report also indicated the role of DFG+1 residue in differentiating Ser versus Thr residues(Chen et al., 2014). These observations hinted to an involvement of Y169 in substrate selectivity. Indeed, our experimental data suggests a pivotal role of Y169 in phosphorylating S32 of IκBα both in vitro and ex vivo (Figure 4D & 4E). It is worthwhile to note that, a phosphoablative mutation at S32 in IκBα to Ile (S32I) is associated with ectodermal dysplasia and T-cell immunodeficiency in patients(Courtois et al., 2003; Mooster et al., 2015). More interestingly, the placement of Y169 in the kinase active site makes it amenable to a phosphorylation-dephosphorylation cycle, suggesting its possible role both in phosphotransfer and substrate specificity. Mutation of Y188, a conserved tyrosine, to non-phosphorylatable phenylalanine abolished the kinase activity of IKK2 in vitro as well as in MEFs upon TNF-treatment (Figure 4C & 4E). The placement of Y188 indicates that its mutation will be detrimental to the structural fold of kinase, thus it is unclear if or how Y188 plays a role in the phosphotransfer. Since phosphotransfer to S32, S36 of IκBα is observed with auto-phosphorylated IKK2 EE, transferred phosphates are not from S177 or S181. Kinase domain autophosphorylation is known to regulate function of other kinases as well, e.g., autophosphorylation at T220 greatly influenced both the activity and substrate specificity of CK1(Cullati et al., 2022).

Removed in the revised version: The strength of input signal, level of IκBα degradation, and NF-κB activation correlate with the amplitude, duration, and kinetics of IKK activity(Behar and Hoffmann, 2013; Cheong et al., 2006). This digital (all-or-none) activation profile of NF-κB appears to be due to rapid activation and inactivation of IKK2. IKK targets multiple substrates in vivo; however, the rapid activation of NF-κB is regulated primarily by IκBα. How IKK specifically selects IκBα for phosphorylation is unclear, although NEMO is reported to direct the kinase activity of IKK2 specifically towards IκBα(Schröfelbauer et al., 2012). We observe that P-IKK2 could transfer phosphate(s) specifically to the S32/S36 of IκBα, and not to other sites (such as the PEST sequence of IκBα) (Figure 6B, 6E and 7A). We surmise if transiency of multisite phosphorylated IKK2 (P-IKK2) underlie this activation-inactivation events during signal response. It is possible that regulatory system of IKK2 is somewhat analogous to the two-component histidine kinase-effector systems in prokaryotes that show a threshold activation response(Bhate et al., 2015; Lamarche et al., 2008; Laub and Goulian, 2007). It is noteworthy that, transfer of a phosphate moiety directly from pHis is energetically much more probable than that from a pTyr, unless it is an enzymatic process which is likely the case for IKK2:IκBα system.

Earlier version: Tyrosines in activation loop of IKK2 are previously reported to undergo signal-induced phosphorylation in cells. These modifications were proposed to be mediated by a tyrosine kinase although there is no evidence of a specific tyrosine kinase(Meyer et al., 2013; Otero et al., 2008; Rieke et al., 2011). Several members of the Ser/Thr kinase family undergo tyrosine phosphorylation and these kinases employ diverse mechanisms for auto-phosphorylation(Bhattacharyya et al., 2006; Ge et al., 2002; Lochhead, 2009; Lochhead et al., 2006, 2005; Tigno-Aranjuez et al., 2010).

Revised version: Tyrosine residues within the activation loop of IKK2 have been previously reported to undergo signal-induced phosphorylation in cells. Although these modifications were proposed to be mediated by a tyrosine kinase, no specific tyrosine kinase has been identified to date(Meyer et al., 2013; Otero et al., 2008; Rieke et al., 2011). Notably, several members of the Ser/Thr kinase family are known to undergo tyrosine phosphorylation, employing diverse mechanisms for autophosphorylation (Bhattacharyya et al., 2006; Ge et al., 2002; Lochhead, 2009; Lochhead et al., 2006, 2005; Tigno-Aranjuez et al., 2010). Regardless of the mechanism responsible for Tyr169 phosphorylation, the phosphate attached to Tyr169 is uniquely primed for transfer to Ser32 of IκBα.

Added in the revised version: Allostery in phosphorylation

Added in the revised version: Several other residues, including tyrosines distributed throughout the kinase molecule, also undergo phosphorylation. While most of these phosphorylations play a role in global stabilization of the kinase, Tyr188 plays a critical role in substrate phosphorylation. Mutation of Tyr188—a conserved tyrosine residue—to a non-phosphorylatable phenylalanine completely abolished IKK2 kinase activity both in vitro and in MEF cells following TNF treatment (Figure 3E & G). The structural placement of Tyr188 suggests that its mutation may allosterically disrupt kinase activity, making its precise role in phosphotransfer unclear. Autophosphorylation of kinase domains is a well-documented regulatory mechanism in other kinases as well. For example, autophosphorylation at Thr220 significantly impacts both the activity and substrate specificity of CK1 (Cullati et al., 2022).

Added in the revised version: Signaling specificity and digital activation

Added in the revised version: The strength of the input signal, along with the extent of IκBα degradation and NF-κB activation, is closely associated with the amplitude, duration, and kinetics of IKK activity (Behar and Hoffmann, 2013; Cheong et al., 2006). The digital (all-or-none) activation profile of NF-κB appears to arise from the rapid activation and inactivation dynamics of IKK2. While IKK phosphorylates multiple substrates in vivo, the swift activation of NF-κB is predominantly governed by IκBα, with IκBβ and IκBε playing lesser roles. This digital activation profile may be intrinsically linked to the phosphate relay process. We propose that Tyr169 serves as both a selection mechanism and a mediator of the rapid phosphorylation kinetics.

Added in the revised version: Interestingly, the regulatory system of IKK2 may be somewhat analogous to the two-component histidine kinase-effector systems in prokaryotes, which exhibit threshold activation responses (Bhate et al., 2015; Lamarche et al., 2008; Laub and Goulian, 2007). It is noteworthy that phosphate transfer directly from pHis is energetically far more favorable than from pTyr, except in enzymatic processes. However, Tyr169 might be special case, behaving similar to active site His in histidine kinases.

Added in the revised version: The conserved salt bridge is specificity determinant not activity

Added in the revised version: Our study highlights the role of K44 in IKK2 autophosphorylation, which is crucial for the specific phosphorylation of Ser32 and Ser36 of IκBα. Interestingly, while the K44 mutant retains non-specific kinase activity, its phosphorylation specificity is compromised. Notably, a monomeric form of IKK2, lacking the NBD and a large portion of the dimeric interface in the SDD, maintains specificity for the N-terminal serines of IκBα (Hauenstein et al., 2014). However, a truncated construct containing only the KD and ULD loses this specificity, behaving like CK2 and phosphorylating the PEST domain (Shaul et al., 2008). Further studies are needed to determine whether a specific structural state of IKK2, possibly influenced by alternative modifications or interactions with other proteins, enables non-specific substrate phosphorylation.

Added in the revised version: Conclusion

Removed in the revised version: The phosphotransfer from P-IKK2 to IκBα monitored here is detection of a single turnover event in absence of the phosphate donor, ATP in the bulk (Figure 6B). Phosphorylation on IKK2 occurs at multiple sites, and not every phosphorylated residue in P-IKK2 expected to transfer its phosphate to IκBα. There is more than one type of phosphorylated residues in P-IKK2 distributed over the length of the kinase, arguably with differential stability wherein only a fraction of the total pool can be relayed. In addition, it is likely that the phosphorylated IKK2 in the bulk harbors a non-homogenous pool where differently phosphorylated species of the kinase are present, and are indistinguishable in our assay schemes. Another intriguing feature of phosphorelay in IKK2 is IKK2’s dependence on ADP (Figure 6B & 6E). It was also observed that, the generic kinase inhibitor AMPPNP (unhydrolyzable ATP analogue) failed to inhibit phosphor-transfer from IKK2 (Figure 7A) and IκBα phosphorylation (Figure S2C). Taken together, we surmise that presence of ADP or AMP-PNP help IKK2 attain a conformational state poised for phosphorylating S32/36 of IκBα even in absence of fresh exogenous ATP. Similar phenomenon was reported for IRE1, where ADP was found to be a more potent stimulator of its ribonuclease activity and that AMPPNP could also potentiate this activity in IRE1(Lee et al., 2008; Sidrauski and Walter, 1997).The list of Ser/Thr kinase family proteins that undergo tyrosine phosphorylation is growing that employ diverse mechanisms for auto-phosphorylation(Bhattacharyya et al., 2006; Ge et al., 2002; Lochhead, 2009; Lochhead et al., 2006, 2005; Tigno-Aranjuez et al., 2010). The most notable example of this dual specificity are DYRK and GSK3β. However, the role of Tyr-phosphorylation in DYRK/GSK3β is not related to that in IKK2, nor does it render DYRKs/ GSK3β capable of relaying phosphates to the substrates. To the best of our knowledge the only other reported example of an EPK, though atypical, transferring phosphate through a phosphoenzyme intermediate is MHCK of Dictyostelium discoideum(Ye et al., 2010).

Removed in the revised version: Our study reveals tyrosine auto-phosphorylation in IKK2 and indicates that consequences of Ser and Tyr autophosphorylations in dual-specific IKK2 are distinct. Moreover, the crucial tyrosine residues are positioned in such a manner that some that are amenable to auto-phosphorylation are essential for substrate phosphorylation. Most surprisingly, the tyrosine-linked phosphate groups can be transferred by a relay mechanism to specific serines in IκBα, an event reminiscing phospho-relay in two-component systems prevalent in prokaryotes and some lower eukaryotes. Our study also reveals the role of K44 of IKK2 in auto-phosphorylation and thus in specific phosphorylation of S32, S36 of IκBα, although the mutant kinase retains its non-specific kinase activity. It may be noted that, while a monomeric version of IKK2 (devoid of NBD and large portion of the dimeric interface in the SDD) still retained its specificity for the N-terminal serines of IκBα(Hauenstein et al., 2014), while a shorter construct containing only the KD and ULD lacked this specificity that phosphorylated the PEST domain like CK2(Shaul et al., 2008). It requires further investigation to determine if a structural state of IKK2, perhaps through a different modification or through binding to another protein, allow it to phosphorylate substrates relatively non-specifically.

Earlier version: In all, it is intriguing how the gateway to the activation of NF-κB is regulated through unique phosphorylation events in IKK2 (Figure 7C). It appears that a multi-layered fail-safe mechanism is devised to tightly control NF-κB activation. Firstly, in absence of upstream signals IKK-complex is unable to phosphorylate Ser 32,36 of IκBα efficiently thereby keeping check on the aberrant and untimely activation of NF-κB. Upon encountering proinflammatory cues, cells need to activate NF-κB immediately and efficiently, i.e., IKK needs to be activated and phosphorylate Ser 32,36 of IκBα. NEMO warrants IKK-activation and ensures that the IKK-complex specifically chooses IκBα from a large pool of substrates of IKK2. IKK2 is designed in such a manner that it phosphorylates itself at activation loop-tyrosines when activated, such that phosphate group(s) can be relayed directly to Ser32,36 of IκBα with great fidelity, thus leaving little chance of a misconstrued signaling event thereby confirming NF-κB activation (Figure 7C). Our discovery might also present the beginning of a new aspect in eukaryotic cell-signaling by EPKs and provide a foundation for mechanistic studies of an intriguing phospho-transfer reaction that helps accomplish the desired specificity. To the best of our knowledge, it is the first report, apart from MHCK of Dictyostelium discoideum (an atypical EPK), to suggest that a metazoan EPK confirms its specificity through the formation of a phosphoenzyme intermediate.

Revised version: In all, it is intriguing how the gateway to the activation of NF-κB is regulated through unique phosphorylation events in IKK2. It appears that a multi-layered fail-safe mechanism is devised to tightly control NF-κB activation. In absence of upstream signals, the IKK-complex is unable to phosphorylate Ser 32,36 of IκBα efficiently thereby keeping check on the aberrant and untimely activation of NF-κB. Upon encountering proinflammatory cues, cells need to activate NF-κB immediately, efficiently and specifically, i.e., IKK needs to be activated and phosphorylate Ser 32,36 of IκBα. NEMO warrants IKK-activation and ensures that the IKK-complex specifically chooses IκBα from a large pool of substrates of IKK2. IKK2 is designed in such a manner that it phosphorylates itself at activation loop-tyrosine when activated, such that phosphate group(s) can be relayed directly to Ser32,36 of IκBα with great fidelity, thus leaving little chance of a misconstrued signaling event thereby confirming NF-κB activation. Our discovery might also present the beginning of a new aspect in eukaryotic cell-signaling by EPKs and provide a foundation for mechanistic studies of an intriguing phospho-transfer reaction that helps accomplish the desired specificity.

Methods summary

Protein purification, auto-phosphorylation assay, immunoprecipitation coupled kinase assay, phosphotransfer assay and other experiments have been described in details in the supplementary document.

Removed in the revised version: Supporting information